```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

# Lab 04: Looping in Parallel

In **Activity 13: The Tabula Recta**, you saw how you could use a collection of characters as a key to work with the Tabula Recta. This is different than the keys you've work with so far, which have been an integer (or integers). As a result, you'll need to determine how to start using strings as keys alongside strings as messages.

Suppose you have two lists of equal length. One contains student names and the other contains their grade on a recent Quick Check quiz.

```python
names = ['Alexis', 'Dee', 'Daryl', 'Sebastian', 'Ora']
grades = [2, 3, 2, 3, 1]
```

Suppose you want to generate a report that prints out a student's name with their grade next to it. If you were trying to avoid a loop, you could do that the **long** way by typing:

```python
print(names[0], grades[0])
print(names[1], grades[1])
print(names[2], grades[2])
print(names[3], grades[3])
print(names[4], grades[4])
```

But that seems a bit repetitive! 


<div class="alert alert-info" role="alert" style="color: black;">

#### Hint

Did you see how if you give the `print` function 2 arguments, separated by commas, it will print them side by side, with just a space in between, EVEN if they are different data types? Nice!
</div>

## Question 1

Inside the function named `q1`, write a `for` loop prints the items contained in `input1` and `input2` in the same manner as above. For now, you can assume that both `input1` and `input2` are lists, and that they have the same number of elements in them.

**Examples:**
```python
>>> q1(names, grades)
Alexis 2
Dee 3
Daryl 2
Sebastian 3
Ora 1

>>> q1([0, 1, 2], [3, 4, 5])
0 3
1 4
2 5

```
<div class="alert alert-info" role="alert" style="color: black;">

#### Hint

Think about how you can compute how many items are in the lists, and then how to use that value to create a `range` of values that would correspond to the indices. Since the lists are the same length, any index that works for one is guaranteed to work on the other.
</div>

```python
def q1(input1, input2):
```

```python
grader.check("q1")
```

## Question 2

Strings and lists aren't that different as they each are a collection of smaller objects, but a string is a collection specifically of individual characters. In the last question you looped over two lists of equal length and printed the element found at each index. 

Write a function named `q2` that prints corresponding characters in two different strings side by side, very similar to the question asked in question 1 that used lists.

**Examples:**
```python
>>> q2('secret', 'blaise')
s b
e l
c a
r i
e s
t e

>>> q2('ncssm', 'smath')
n s
c m
s a
s t
m h
```

Remember, you should be writing a function with a loop that works no matter what characters the strings contain.

```python
def q2(string1, string2):
```

```python
grader.check("q2")
```

## Using Indices With Strings

You likely have already realized that this question is easiest solved by not iterating character by character over each string, but rather index value by index value. By writing a loop that starts:

`for i in range(len(message))`

You can use `i`, the string index to reference corresponding characters in *any* string, in this case both `keyword` and `message`. If you had started with a loop like:

`for char in message:`

you would have had an easy way to pull each character out of the `message` string, but no easy way to match it up with the corresponding character in `keyword`. For parallel looping, indices are your friend!

## Strings of Different Lengths

So, what can you do when strings have unequal length? The Tabula Recta method assumes that your keyword is equally long as the message, so each character in the message has a corresponding character in the key to work with. 

Keys are usually decided upon ahead of time, but the length of messages will vary from message to message. We'll need a way to make sure that any given key is properly used with variable length messages. 

We saw in the lesson that you could just repeat the key until it's just long enough to **"cover"** the message. To do that in Python we have two ways to mimic this step:

1. Take your key as a string and repeatedly append copies of itself to the end until it's long enough to match the message.
2. Use a mod operator `%` to ensure that the correct index from the key is referenced , even when index values for the message get to be larger than the index values for the key.

Let's explore both options so you can see which you prefer.


### Question 3

In the cell below, write a function that takes in a Tabula Recta key (which is a string) and a message (which is also a string) and computes how many times the keyword needs to be repeated to "cover" the entire message (meaning, every character in the message has a corresponding character in the key. Then, use this value to **return** a string named `covered_key` that perfectly "covers" the message.

**Examples:**

```python
>>> q3('ABCDEFG', 'thisisalongermessage')
'ABCDEFGABCDEFGABCDEF'

>>> q3('1234', 'wowthismessageisalsoprettylong')
'123412341234123412341234123412'

>>> q3('MATH', 'somehowthismessageisevenlongerthanbothoftheothermessages')
'MATHMATHMATHMATHMATHMATHMATHMATHMATHMATHMATHMATHMATHMATH'

```python
def q3(key, message):
    key_length = ...
    message_length = ...
    ...
```

```python
grader.check("q3")
```

## Question 1.4

The other option to work with keys and messages of unequal length is to use the mod operator `%` on the index for the message to compute the index of the corresponding character in the key. 

For example, if you are currently working with at the character at index `107` in your message, and your original key was 5 characters long, which character in the keyword should be used to encipher the message? You should be able to compute that value using the current index `107`, and the length of your key, `5`.

Write a function named `q4` that computes the index of the corresponding character in `key` that *would* be lined up at `i`, which represents the index for a particular message.

**Examples:**
```python
>>> q4('ABCDEFG', 107)
'C'

>>> q4('SECRETKEY', 1953)
'S'

>>> q4('PLEASEDONTGUESS', 323423535234)
'T'
```

```python
def q4(key, i):
```

```python
grader.check("q4")
```

# Submitting your work
You're done with this Lab! All assignments in the course will be distributed as notebooks like this one, and you will submit your work by doing the following:
* Save your notebook
* Restart the kernel and run up to this cell.
* Run all the tests by running the cell containing `grader.check_all()`. Make sure they pass the way you expect them to.
* Run the cell below with the code `grader.export(...)`.
* Download the file named `labXX.zip`, found in the explorer pane on the left side of the screen.
* Upload `labXX-<date-time stamp>.zip` to the corresponding lab assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1 (points: 3)
```
>>> q1(names, grades)
Alexis 2
Dee 3
Daryl 2
Sebastian 3
Ora 1
```
```
>>> q1([0, 1, 2], [3, 4, 5])
0 3
1 4
2 5
```
```
>>> q1(range(0, 10), range(10, 20))
0 10
1 11
2 12
3 13
4 14
5 15
6 16
7 17
8 18
9 19
```


Tests for q2 (points: 3)
```
>>> q2('secret', 'blaise')
s b
e l
c a
r i
e s
t e
```
```
>>> q2('ncssm', 'smath')
n s
c m
s a
s t
m h
```
```
>>> q2('hello', 'world')
h w
e o
l r
l l
o d
```


Tests for q3 (points: 3)
```
>>> q3('ABCDEFG', 'thisisalongermessage')
'ABCDEFGABCDEFGABCDEF'
```
```
>>> q3('1234', 'wowthismessageisalsoprettylong')
'123412341234123412341234123412'
```
```
>>> q3('MATH', 'somehowthismessageisevenlongerthanbothoftheothermessages')
'MATHMATHMATHMATHMATHMATHMATHMATHMATHMATHMATHMATHMATHMATH'
```


Tests for q4 (points: 3)
```
>>> q4('ABCDEFG', 107)
'C'
```
```
>>> q4('SECRETKEY', 1953)
'S'
```
```
>>> q4('PLEASEDONTGUESS', 323423535234)
'T'
```
