```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03.ipynb")
```

# Homework 03: Public key encryption

This homework will focus on topics covered in the final week of the course: public key cryptography.

## Helper functions

To get you started, run the cell below to load functions that have been written in earlier homework assignments. Functions that are included:
* `gcd(<a>,<b>)`. This function will compute the greatest common divisor between `a` and `b`
* `egcd(<a>, <b>)`. Performs the Extended Euclidean Algorithm on inputs `a` and `b`
* `mod_inverse(<e>, <n>)`. Computes the multiplicative inverse of `e` in modulus `n`
* `phi(<n>)`. This is Euler's totient function as described in [11.3 RSA](https://macs4200.org/chapters/11/3/rsa.html)

You can use these functions in any of your code below after you've imported them.

```python
from helpers import gcd, egcd, mod_inverse, phi
```

## 1. KidRSA
*Recommended Reading:* https://www.macs4200.org/chapters/11/1/kidrsa.html

In this question you will implement the KidRSA algorithm.

### Part A: Key Generation

Write a function `kidrsa_keygen` that uses the values for `a`, `b`, `a_prime`, and `b_prime` to compute a pair of public and private keys using the Kid RSA key generation algorithm. The function will return the pair of keys as a list of lists. The first key in the list should be the public key $\left(e,n\right)$ and the second key in the list should be the private key $\left(d,n\right)$.

**Sample Test Cases**:
```python
>>> kidrsa_keygen(3,7,5,11)
[[103, 1169], [227, 1169]]

>>> kidrsa_keygen(2,324,100,8)
[[64702, 550017], [5500, 550017]]
```

```python
def kidrsa_keygen(a, b, a_prime, b_prime):
    
    ...
```

```python
grader.check("q1_a")
```

### Part B: Encryption / Decryption

Write a function `kidrsa` that takes in a numerical integer message, a list containing either a public or private key, and encrypts or decrypts the message. Since the process for encryption and decryption are identical, you do not need a boolean `decrypt` parameter as in earlier homework assignments.

`kidrsa` should output a numerical integer that represents ciphertext when the function is provided numerical plaintext, and vice-versa.

**Reminder:** Your message, `message`, must be less than the modulus `n` from the key **and** be relatively prime to `n` (i.e `gcd(m,n) == 1`). Your function **must** check for this condition and return the boolean value `False` if it does not meet these criteria.

**Sample Test Cases**:
```python
>>> kidrsa( 348, [18221, 347273])
89994

>>> kidrsa( 89994, [9377, 347273])
348
```

If `message > n`
```python
>>> kidrsa( 500000, [18221, 347273])
False
```

If `gcd(message, n) != ``
```python
>>> kidrsa( 61, [18221, 347273])
False
```

```python
def kidrsa(message, key):
    
    ...
```

```python
grader.check("q1_b")
```

### Part C: Cracking Keys

KidRSA is not secure because it's fairly easy and fast to determine a private key if you have someone's private key. Write a function `crack_kidrsa` that accepts a public key (as a list) as the input, and returns the private key.

`crack_kidrsa` must return an a list that represents the private key formatted as `[d, n]`. `crack_kidrsa` must check to ensure that the provided value of `e` has an inverse in mod `n`. If they aren't the function should return the boolean value `False`

**HINT:** You should not brute force this! You should use the `mod_inverse` function to help here. `mod_inverse` returns the boolean value `False` if there is no multiplicative inverse found. 

**Sample Test Cases**:
```python
>>> crack_kidrsa( [18221, 347273] )
[9377, 347273]

>>> crack_kidrsa( [4, 12] )
False
```

```python
def crack_kidrsa(public_key):
    
    ...
```

```python
grader.check("q1_c")
```

## 2. The RSA Cipher
*Recommended Reading:* https://www.macs4200.org/chapters/11/3/rsa.html

### Part A: Key Generation

Write a function `rsa_keygen` that uses the prime integers `p` and `q`, and the encryption exponent `e` to compute a pair of public and private keys using the RSA key generation algorithm. `e` should take on the default value of `65537`, but should allow for optional custom values of `e` to be provided as an input.

The function must return the pair of keys as a list of lists. The first key in the list should be the public key $\left(e,n\right)$ and the second key in the list should be the private key $\left(d,n\right)$.

`rsa_keygen` must check that `e` is a valid encryption exponent. If it's not, `rsa_keygen` should return the boolean value `False`

**Reminder:**  $e$ and $\varphi(n)$ need to be relatively prime for $e$ to be valid.

**Sample Test Cases**:
```python
>>> rsa_keygen(3, 11, 3)
[[3, 33], [7, 33]]

>>> rsa_keygen(1787, 3881)
[[65537, 6935347], [5655233, 6935347]]

>>> rsa_keygen(1787, 3881, e=36)
False
```

```python
def rsa_keygen(p, q, e=65537):
    
    ...
```

```python
grader.check("q2a")
```

### Part B: Encryption / Decryption

Write a function `rsa` that takes in a numerical integer message, a list containing either a public or private key, and encrypts or decrypts the message. Since the process for encryption and decryption are identical, you do not need a boolean `encipher` parameter as in earlier homework assignments.

`rsa` should output a numerical integer that represents ciphertext when the function is provided numerical plaintext, and vice-versa.

**Reminder:** Your message, `message`, must be less than the modulus `n` from the key. Your function must check for this condition and return the boolean value `False` if it does not meet the criteria.

**Hint:** You should use the `pow` function instead of the `**` operator so you're using the fast modular exponentiation algorithm. This will drastically reduce the time to encrypt or decrypt a message. The numbers in this example are far too large for Python to efficiently carry out the exponentiation with the `**` operator!

**Sample Test Cases**:
```
>>> rsa( 5, [3,33] )
26

>>> rsa( 26, (7,33) )
5

>>> rsa( 234712, [65537, 6935347])
2678654

>>> rsa( 2678654, [5655233, 6935347])
234712

m > n case
>>> rsa( 26786543243423, [5655233, 6935347])
False
```

```python
def rsa(message, key):
    
    ...
```

```python
grader.check("q2_b")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1_a (points: 2)
```
>>> kidrsa_keygen(3, 7, 5, 11) != None
True
```
failure_message: Make sure you are using a return statement in your function, not just printing the output
```
>>> type(kidrsa_keygen(3, 7, 5, 11)[0][1]) != float
True
```
failure_message: The value for n should be an integer, not a float. HINT dividing with // returns an integer.
```
>>> kidrsa_keygen(3, 7, 5, 11)
[[103, 1169], [227, 1169]]
```
failure_message: Your keys are incorrect
```
>>> kidrsa_keygen(2, 324, 100, 8)
[[64702, 550017], [5500, 550017]]
```
failure_message: Your keys are incorrect


Tests for q1_b (points: 3)
```
>>> kidrsa(348, [18221, 347273]) != None
True
```
failure_message: Make sure you are using a return statement in your function, not just printing the output
```
>>> kidrsa(348, [18221, 347273])
89994
```
failure_message: Sample case not working
```
>>> kidrsa(89994, [9377, 347273])
348
```
failure_message: Sample case not working
```
>>> kidrsa(500000, [18221, 347273])
False
```
failure_message: Sample case not working
```
>>> kidrsa(61, [18221, 347273])
False
```
failure_message: Sample case not working


Tests for q1_c (points: 3)
```
>>> crack_kidrsa([18221, 347273]) != None
True
```
failure_message: Make sure you are using a return statement in your function, not just printing the output
```
>>> crack_kidrsa([18221, 347273])
[9377, 347273]
```
failure_message: Sample case not working
```
>>> crack_kidrsa([4, 12])
False
```
failure_message: Sample case not working


Tests for q2_b (points: 3)
```
>>> rsa(5, [3, 33]) != None
True
```
failure_message: Make sure you are using a return statement in your function, not just printing the output
```
>>> rsa(5, [3, 33])
26
```
failure_message: Are you cleaning your message?
```
>>> rsa(234712, (65537, 6935347))
2678654
```
failure_message: Is your deciphered text lowercase?
```
>>> rsa(2678654, (5655233, 6935347))
234712
```
failure_message: Is your ciphertext uppercase?
```
>>> rsa(26786543243423, (5655233, 6935347))
False
```
failure_message: Is your ciphertext blocked?


Tests for q2a (points: 2)
```
>>> rsa_keygen(3, 11, 3) != None
True
```
failure_message: Make sure you are using a return statement in your function, not just printing the output
```
>>> rsa_keygen(1787, 3881)
[[65537, 6935347], [5655233, 6935347]]
```
failure_message: Sample case not working
```
>>> rsa_keygen(1787, 3881, e=36)
False
```
failure_message: Sample case not working
