```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

# Homework 01: Monoalphabetic Ciphers and Cryptanalysis

## Overview
This homework will cover the topics taught in Lessons 01 - 12.

## Recommended Textbook Readings:
* [Loops](https://macs4200.org/chapters/03/3/intro-to-loops)
* [Functions](https://macs4200.org/chapters/03/4/functions)
* [Chi Squared Scoring](https://macs4200.org/chapters/06/4/chi-squared-scoring)
* [Affine Cipher](https://macs4200.org/chapters/04/6/affine-cipher)

<div class="alert alert-warning" role="alert"  style="color: black;">
    
### Allowable resources

You may use any programming techniques covered in class to answer the following questions on this programming portion of this assignment, however, pay careful attention to the directions for each question to ensure you are using the specific approach requested, when applicable. You are only allowed to use your notes and refer to previous lesson materials and activities as resources to complete this assignment. No artificial intelligence tools, textbooks other than the book for this course, or internet resources may be used to complete this programming portion of the assignment. This is an individual assignment, so you may not obtain help from any other students to complete the assignment.
</div>

<div class="alert alert-danger" role="alert"  style="color: black;">
    
### This homework has hidden tests in it!

Additional tests will be run once your homework is submitted for grading. While you may pass all the tests you have access to before submission, you may not earn full credit if you do not pass the hidden tests as well.

Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and/or you have given an answer that *could* make sense in context. For example, a test you have access to while completing the assignment may check that you selected a valid choice for a multiple choice problem (1, 2, or 3) or that your answer is an float between 0 and 1 if asked to compute a proportion. The tests that are run after submission will further evaluate your work, typically to check for accuracy. 

**Do not assume that just because all your tests pass before submission means that your answers are fully correct!** You should test your functions thoroughly to ensure they meet the requirements listed in the problem statement.
</div>

## Part 1: Perfect Numbers

A positive integer $n$ is called a **perfect number** if all divisors of $n$ which are less than $n$ sum to equal $n$. 

Put another way, a perfect number is a positive integer that is equal to the sum of its positive factors, excluding the number itself. For example, $6$ is a perfect number because the factors of $6$ that are less than $6$ are $1$, $2$, and $3$, which sum to $6$ ($6 = 1 + 2 + 3$).

### Question 1.1

Write a function named `perfect` which accepts an integer `n` as an input. The function `perfect` should return `True` if the integer `n` is a perfect number, and return `False` if integer `n` is not a perfect number.

For example:

```python
>>> perfect(6)
True

>>> perfect(28)
True

>>> perfect(15)
False
```

This function should return `False` if it is called with an non-postive integer as an input:


```python
>>> perfect(-2)
False

>>> perfect(0)
False
```


<div class="alert alert-info" role="alert" style="color: black;">

#### Hint

You'll need to use the mod operator `%` to help with this question!

</div>

```python
def perfect(n):
    ...
```

```python
grader.check("q1_1")
```

### Question 1.2
Write a function `all_perfects_less_than` which accepts any integer `k` and returns a **list** of all perfect numbers less than `k`.

For example:
```python
>>> all_perfects_less_than(50)
[6, 28]

>>> all_perfects_less_than(500)
[6, 28, 496]

>>> all_perfects_less_than(5)
[]
```

<div class="alert alert-info" role="alert" style="color: black;">

#### Hint

Make sure your function works with positive, negative, and 0 integer values.

</div>

```python
def all_perfects_less_than(k):
    ...
```

```python
grader.check("q1_2")
```

## Part 2: Bar Charts and Chi Squared scoring

In this part, you'll be using chi squared scoring to perform some cryptanalysis. Run the code cell below to load in the `caesar` function you've written previously written in Lab 02.


<div class="alert alert-info" role="alert" style="color: black;">

#### Note
In this homework, it's important that you use the provided functions, not functions you've written in class, to ensure that your results are consistent with the automatic grader checks.
</div>

```python
from helpers import caesar
```

You can now use `caesar` as if it were defined in this notebook, which you'll need to complete some questions in Part 2.

```python
caesar('HELLO', 3)
```

<!-- BEGIN QUESTION -->

## Question 2.1

Complete the following function `chi_squared` so it returns the chi-squared score for a provided piece of input text, `text`, based on the provided alphabet in `LETTERS`.

For example:

```python
>>> chi_squared('HELLO')
22.526108340863132
```

```python
>>> chi_squared('BERSJDFYUIBSDF')
85.74576834801644
```

```python
def chi_squared(text, LETTERS='ABCDEFGHIJKLMNOPQRSTUVWXYZ'):

    letter_proportion =[0.082, 0.015, 0.028, 0.043, 0.127, 0.022, 0.02, 0.061, 0.07, 0.0015, 0.0077, 0.04, 0.024, 0.067, 0.075, 0.019, 0.00095, 0.06, 0.063, 0.091, 0.028, 0.0098, 0.024, 0.0015, 0.02, 0.00074]

    ...
```

```python
grader.check("q2_1")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2.2

You intercepted the following ciphertext that you suspect was encrypted using the Caesar cipher. Notice that this ciphertext is already cleaned for analysis.

```python
ciphertext = 'LZWVWTSLWLGMJFSEWFLLZWTMKVWHSJLWVSLXANWLZAJLQSEUSJJQAFYGMJVWTSLWLWSELZJGMYZLZWVSJCFWKKLGOSJVLZWKLSLWUZSEHAGFKZAHWAYZLTDWSJQWQWVKLMVWFLKUDMLUZWVLJSNWDEMYKGXUGXXWWSFVJWNAWOWVFGLWUSJVKTQLZWVAEGNWJZWSVDAYZLKJWEWETWJALKFGLBMKLSTGMLXSULKUGSUZDSOKGFKSAVXJGELZWXJGFLKWSLALKSTGMLUGFFWULAFYOALZLZWBMVYWKESCWLZWEUSJWSTGMLQGMJSJYMEWFLAFGVVWVLZGMYZEQKLGESUZUZMJFWVOALZSFPAWLQLZAKOSKEQXAJKLLAEWUGEHWLAFYSLLZWKLSLWDWNWDLZWHJWNAGMKFAYZLAVHJSULAUWVEQGHWFAFYKLSLWEWFLMFLADEQJGGEESLWLZJWSLWFWVLGKEGLZWJEWOALZSHADDGOAXAVAVFLDWLZWJKDWWHLZWUGEHWLALAGFOSKZWDVSLUWFLJSDMFANWJKALQALKAEHGKAFYTJAUCTMADVAFYKESCAFYEWXWWDKESDDWJLZSFMKMSDLWSEKXJGESUJGKKLZWKLSLWZMVVDWVAFZSDDOSQKGJUDSAEWVLSTDWKAFLZWKLMVWFLMFAGFSTMRRAFYZANWGXFWJNGMKWFWJYQSFVDSKLEAFMLWHJWHSJSLAGFKGMJXAJKLJGMFVLZWLGHAUOSKUDAESLWHGDAUQOWFLKMJHJAKAFYDQOWDDEQHSJLFWJVSEAWFSFVASVNSFUWVLGLZWIMSJLWJXAFSDKLZWFKGEWZGOLGLZWKWEAXAFSDKOALZWSUZJGMFVEQUGFXAVWFUWYJWOLZWZGMJKGXJWKWSJUZLZWHJSULAUWVWTSLWKLZWXWWVTSUCKWKKAGFKALOSKSDDUGEAFYLGYWLZWJLZWXAFSDJGMFVOSKSFAFLWDDWULMSDWPLJSNSYSFRSSYSAFKLLZWVWXWFVAFYUZSEHAGFKXJGEOWKLDSCWSUSVWEQLZWSMVAWFUWZSVYJGOFLGFWSJDQSZMFVJWVHWGHDWHSJWFLKUGSUZWKWDAEAFSLWVUGEHWLALGJKWNWFSXWOMFANWJKALQKLMVWFLKUMJAGMKSTGMLLZWUGEEGLAGFOZWFLZWBMVYWKSFFGMFUWVVSEAWFSFVEWSKLZWOAFFWJKAUGMDVFLHJGUWKKALSLXAJKLUGSUZDSOKGFKWFLZMKASKLAUWETJSUWXAFSDDQESVWALJWSDOWZSVVGFWALGMJKESDDHMTDAUKUZGGDZSVVWXWSLWVHJANSLWSUSVWEAWKOALZVWVAUSLWVVWTSLWUGSUZWKSFVDSNAKZJWKGMJUWKGFLZWJAVWZGEWLJGHZQFWKLDWVKSXWDQAFEQDSHAJWSDARWVKGEWLZAFYAEHGJLSFLTWAFYMFVWJWKLAESLWVUSFTWSHGOWJXMDSVNSFLSYWFGGFWWPHWULWVMKLGOAFWPUWHLUGSUZDSOKGFOZGZSVTWDAWNWVAFMKXJGELZWTWYAFFAFYKGEWLAEWKLZSLKSDDLZWSVNSFLSYWQGMFWWV'
```

In the code cell below, generate a bar chart of the character frequency in the ciphertext. The height of the bar should indicate the *proportion* of the ciphertext that contains the letter associated with each bar (not the raw count).

```python
# Add your code to this cell so that when it is run, a bar chart appears below
# All your code should be contained in this single cell please!
```

<!-- END QUESTION -->

### Question 2.3

In the code cell below, write a block of code that:
1. uses every possible key for the Caesar cipher to decrypt the text,
2. score each decryption using the chi squared method, and
3. store each score to a list named `chi_squared_scores`, in order. 

For example, the first element in the `chi_squared_scores` list, `chi_squared_scores[0]`, should contain the chi squared score for the **decryption** obtained using a key of 0, `chi_squared_scores[1]` should contain the score when using a key of 1, and so on.

```python
chi_squared_scores = []


print(chi_squared_scores)
```

```python
grader.check("q2_3")
```

### Question 2.4

Based on the list of chi squared scores you computed in the previous question, determine which key was most likely used to encrypt the message. 

**Assign the value you think is the key to the variable `likely_key` in the code cell below.**

```python
likely_key = ...
```

```python
grader.check("q2_4")
```

## Part 3: The Affine cipher

The goal in this question is to write a function that implements the affine cipher. To do this, you'll need a few more "helper" functions that you've written previously, or, implement some mathematics that we've discussed previously.

The code cell below loads the following functions:
* `text_clean`. The function we've written before that ensures the input string, `text` only contains characters found in the provided alphabet in the string `LETTERS`.
* `text_block`. The function that groups characters in the input string, `text` into blocks of `n` characters, each block separated by a single space. `n` has a default value of 5.
* `eea`. This function implements the extended euclidean algorithm to determine the modular inverse of input `k` in modulus of input `m`. You do NOT need to understand how this function works, as it requires a bit more programming knowledge than is appropriate for this course.

```python
from helpers import text_clean, text_block, eea
```

Since you can't see the code for these functions, you'll need to know what inputs they are expecting. You can do that using the `help` function, and providing it the name of the function. It prints the function "header" that contains all the argument names, and any default values they take.

```python
help(text_clean)
```

```python
help(text_block)
```

```python
help(eea)
```

Use the code cell below to make sure you know how to use these functions before proceeding. `text_clean` and `text_block` should be familiar since you've written these functions yourself before, so pay close attention to `eea` to make sure you can use it to calculate multiplicative inverses in different modulus.

```python

```

### Question 3.1

Use the provided functions to complete the `affine` function below.

The function `affine` should have arguments:
* `text`: a string that contains a message to encrypt or decrypt. You should assume this string is not "clean"
* `m_key`: an integer that will be used as the multiplicative encryption key
* `a_key`: an integer that will be used as the additive key
* `decrypt` (optional argument set to `False`): a boolean value (e.g. True or False) that determines if the function encrypts a message (`decrypt = False`) or decrypts a message (`decrypt = True`).
* `LETTERS` (optional argument, default value `'ABCDEFGHIJKLMNOPQRSTUVWXYZ'`): the alphabet that is considered valid for the message

The `affine` function should return either:
* an encrypted ciphertext message as an upper-case and blocked string (when `decrypt=False`)
* a decrypted plaintext message as a lowercase string (when `decrypt=True`)

For example, the following function calls should return the following outputs:
```python
>>> affine('hello friends', 15, 4)
'FMNNG BZUMR XO'
```

```python
>>> affine('FMNNG BZUMR XO', 15, 4, decrypt=True)
'hellofriends'
```

```python
>>> affine('hello friends', 15, 4, LETTERS='ABCDEFG')
'BCBA'
```

```python
>>> affine('BCBA', 15, 4, decrypt=True, LETTERS='ABCDEFG')
'efed'
```

<div class="alert alert-info" role="alert" style="color: black;">

#### Hint
Make sure your `affine` function works with custom alphabets assigned to `LETTERS`.

</div>


```python
def affine(text, m_key, a_key, decrypt=False, LETTERS = 'ABCDEFGHIJKLMNOPQRSTUVWXYZ'):
    cleaned_text = ...
    m_key_inverse = ...
    final_text = ''
    
    ...
```

```python
grader.check("q3_1")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1_1 (points: 4)
```
>>> perfect(6) == True
True
```
failure_message: 6 should be perfect
```
>>> perfect(28) == True
True
```
failure_message: 28 should be perfect
```
>>> perfect(496) == True
True
```
failure_message: 496 should be perfect
```
>>> perfect(2) == False
True
```
failure_message: 2 is not perfect
```
>>> perfect(15) == False
True
```
failure_message: 15 is not perfect
```
>>> perfect(1) == False
True
```
failure_message: 1 is not perfect
```
>>> perfect(0) == False
True
```
failure_message: 0 is not perfect
```
>>> perfect(-6) == False
True
```
failure_message: Negative numbers are not perfect


Tests for q1_2 (points: 4)
```
>>> all_perfects_less_than(1) == []
True
```
failure_message: For k < 6, the function should return an empty list
```
>>> all_perfects_less_than(0) == []
True
```
failure_message: For k < 6, the function should return an empty list
```
>>> all_perfects_less_than(2) == []
True
```
failure_message: For k < 6, the function should return an empty list
```
>>> all_perfects_less_than(7) == [6]
True
```
failure_message: For 6 < k <= 28, the only perfect number is 6
```
>>> all_perfects_less_than(28) == [6]
True
```
failure_message: For 6 < k <= 28, the only perfect number is 6
```
>>> all_perfects_less_than(50) == [6, 28]
True
```
failure_message: for k = 50, both 6 and 28 are perfect numbers
```
>>> all_perfects_less_than(500) == [6, 28, 496]
True
```
failure_message: for k = 50, 6, 28, and 496 are perfect numbers
```
>>> all_perfects_less_than(-10) == []
True
```
failure_message: Negative values for k should also return an empty list


Tests for q2_1 (points: 4)
```
>>> isinstance(chi_squared('HELLO'), float)
True
```
failure_message: chi_squared should return a float value
success_message: chi_squared returns a float
```
>>> round(chi_squared('A'), 4)
11.1963
```
failure_message: chi_squared does not compute the correct value for a single character
success_message: chi_squared computes the correct value for a single character
```
>>> round(chi_squared('TOBEORNOTTOBETHATISTHEQUESTIONWHETHERTISNOBLERINTHEMINDTOSUFFERTHESLINGSANDARROWSOFOUTRAGEOUSFORTUNEORTOTAKEARMSAGAINSTASEAOFTROUBLESANDBYOPPOSINGENDTHEM'), 4) < 50
True
```
failure_message: chi_squared does not give the expected score for a longer English passage
success_message: chi_squared gives a low score for a longer English passage
```
>>> round(chi_squared('Z' * 26), 4) > 30000
True
```
failure_message: chi_squared should give a very high score for all Z's
success_message: chi_squared gives a very high score for a repeated rare letter


Tests for q2_3 (points: 2)
```
>>> len(chi_squared_scores) == 26
True
```
failure_message: chi_squared_scores should have 26 elements
success_message: chi_squared_scores has 26 elements
```
>>> import math
>>> math.floor(chi_squared_scores[1]) != 20790 and math.floor(chi_squared_scores[25]) != 19194
True
```
failure_message: make sure you are deciphering the message, not enciphering the message
```
>>> import numpy as np
>>> bool(np.isclose(chi_squared_scores[0], 11476.590345085184))
True
```
failure_message: The first element of chi_squared_scores is incorrect
success_message: The first element of chi_squared_scores is correct
```
>>> import numpy as np
>>> bool(np.isclose(chi_squared_scores[25], 20790.189751111415))
True
```
failure_message: The last element of chi_squared_scores is incorrect
success_message: The last element of chi_squared_scores is correct


Tests for q2_4 (points: 1)
```
>>> type(likely_key) == int
True
```
failure_message: likely_key is not an integer. Did you put the chi-squared score instead?
success_message: likely_key has been assigned to an integer value


Tests for q3_1 (points: 6)
```
>>> affine('hello friends', 15, 4)
'FMNNG BZUMR XO'
```
failure_message: affine does not work for a simple encryption test case
success_message: affine works for a simple encryption test case
```
>>> affine('FMNNG BZUMR XO', 15, 4, decrypt=True)
'hellofriends'
```
failure_message: affine does not work for a simple decryption test case
success_message: affine works for a simple decryption test case
```
>>> affine('hello friends', 15, 4, LETTERS='ABCDEFG')
'BCBA'
```
failure_message: affine does not work for a simple encryption test case with non-standard LETTERS
success_message: affine works for a simple encryption test case with non-standard LETTERS
```
>>> affine('BCBA', 15, 4, decrypt=True, LETTERS='ABCDEFG')
'efed'
```
failure_message: affine does not work for a simple encryption test case with non-standard LETTERS
success_message: affine works for a simple encryption test case with non-standard LETTERS
```
>>> affine('HeLLo', 3, 7)
'CTOOX'
```
failure_message: affine does not handle mixed-case input correctly
success_message: affine handles mixed-case input correctly
```
>>> affine('hello', 3, 1, LETTERS='ABCDEFGHIJ')
'CD'
```
failure_message: affine does not work with a custom LETTERS alphabet
success_message: affine works with a custom LETTERS alphabet
