```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03.ipynb")
```

# Lab 03: Functions and File I/O



## Writing functions
There are several tasks that you'll need to complete frequently throughout this course when working with text. In this lab, you'll learn how to write a function to complete these tasks so you don't need to copy, paste, and modify your computer code each time you want to slightly modify the parameters of the task at hand. You'll write functions to help perform the following tasks:
* `text_clean`: This function will accept a string as an input, and return a string that only contains characters in the alphabet you're using
* `text_block`: Ciphertext is traditionally displayed as upper-case letters grouped into blocks of 5 characters. This function will accept a string and return the same string grouped in blocks of 5 characters with a single space between the blocks
* `caesar`: This function will accept a string and encrypt or decrypt it using the Caesar cipher, and return a string with the result
* `gcd`: This function will accept two integers as inputs and use the Euclidean Algorithm to determine the greatest common divisor, which it will return as an integer

### Question 1

When we use python to encrypt messages, we want our input text to be predictable so we can provide instructions to the computer that we know will work with the provided input. That often means we need to "clean" our plaintext before encryption. We often want to strip out punctuation and spaces, and ensure that all the characters are part of the alphabet we're using.

Write a function named `text_clean` with arguments:
* `text`: a string containing the text to clean
* `LETTERS` (optional argument, default value `'ABCDEFGHIJKLMNOPQRSTUVWXYZ'`): a string that defines the valid alphabet
* `match_case` (optional argument, default value `False`): a boolean that decides if the function should act is if `'A'` and `'a'` should count as the same character or not.

The function should check each character in `text` to see if is in the string `LETTERS`. If it is, the character should be kept in the output of the function. If it is not, it should not be in the output of the function.

Examples:

```python
>>> text_clean('Hello')
"HELLO"
```

```python
>>> text_clean('Hello', match_case=True)
"H"
```

```python
>>> text_clean("Hello, World!")
"HELLOWORLD"
```

```python
>>> text_clean("a1b2c3", LETTERS="0123456789")
"123"
```


**Note**: You should not use the `.replace` string method to remove unwanted characters from `text`. This would require you, the programmer, to construct a *deny-list* of invalid characters. Users of your function will likely be very good at coming up with invalid characters you didn't think of putting on your deny-list. Instead, check to see if the characters in `text` are in `LETTERS`, you're using an *allow-list* approach, which is much easier to use to create a valid string.

```python
# Your code goes here

```

```python
grader.check("q1")
```

### Question 2

Since our text cleaning process usually involves stripping out spaces, our encrypted text will just be one long string of characters without spaces. While this is helpful to disguise word length, it's difficult to read without spaces to break up the text. It's standard practice to "block" our text into groups of characters, typically 4 or 5 characters in a block.

Write a function named `text_block` with arguments:
* `text`: a string containing the text to block. You can assume it's "clean"
* `n` (optional argument, default value `5`): an integer that sets the size of each block

The function should return a string equivalent to `text`, but with a single space inserted every `n` characters. For example:

```python
>>> print(text_block('TESTMESSAGE'))
TESTM ESSAG E
```

```python
>>> print(text_block('TESTMESSAGE'), n=3)
TES TME SSA GE
```

```python
# Your code goes here

```

```python
grader.check("q2")
```

### Question 1.3

Now that you can clean and block messages, you're ready to write a function that implements the Caesar cipher.

Write the function `caesar` with arguments:
* `text`: a string that contains a message to encrypt or decrypt. You should assume this string is not "clean"
* `key`: an integer that will be used as the encryption or decryption key
* `decrypt` (optional argument set to `False`): a boolean value (e.g. True or False) that determines if the function encrypts a message (`decrypt = False`) or decrypts a message (`decrypt = True`).
* `LETTERS` (optional argument, default value `'ABCDEFGHIJKLMNOPQRSTUVWXYZ'`): the alphabet that is considered valid for the message

The `caesar` function should return:
* an encrypted ciphertext message as an upper-case and blocked string
* a decrypted plaintext message as a lowercase string

```python
# Your code goes here

```

```python
grader.check("q3")
```

### Question 4

Now, write a function that computes the greatest common divisor between two numbers using the Euclidean Algorithm. 

Write the function named `gcd` with arguments:
* `a`: an integer
* `b`: an integer

which returns an integer that represents the greatest common divisor of `a` and `b`.

```python
# Your code goes here

```

```python
grader.check("q4")
```

Now that you have a bunch of functions to help you out, let's get to work applying them to more interesting questions.

## File Input / Output

In Python, you can read in text from a simple text file that's generated by word processors. These types of files end in `.txt`. `.txt` files are not fancy; you can not have different fonts, font sizes, colors, images, etc. These files *only* contain characters. Most word processors like Microsot Word and Google Docs have an option to save your documents as `.txt` files, but you can also use Notepad (on Windows) or Text Editor (Mac) to easily create these files as well.

You can load the text from a `.txt` file as a string in Python, and assign it to a variable. Alongside this lab notebook you'll find a file named `caesar-ciphertext.txt`. Here's the command to read it in as a string named `caesar_ciphertext`.

```python
with open('caesar-ciphertext.txt') as f:
    caesar_ciphertext = f.read()

print(caesar_ciphertext)
```

You can see that this is an encrypted message that has only capital letters and spaces to group the characters into blocks of length 5. We'll attempt to decipher this message.

### Question 5

In the cell below, use your `caesar` function from earlier in this assignment to decrypt the message, and assign it to the variable named `plaintext`.

**Note:** You don't have the key, so you'll need to try a few to figure it out.

```python
plaintext = ...
plaintext
```

```python
grader.check("q5")
```

### Question 6

You can also save a string to a text file, allowing you to send it by email, upload the file to Google Drive, etc. The code to save a string stored to a text file, which must end in `.txt`, is discussed in [Chapter 5, Section 3](https://macs4200.org/chapters/05/3/file-io#saving-text-to-file) of our textbook. Adapt the code shown in the textbook to save the string stored to `plaintext` to the file named `caesar-plaintext.txt`.

Once you've updated the code, run the cell below to save your plaintext message as a `.txt` file.

```python
# Your code goes here

```

```python
grader.check("q6")
```

You should see `caesar-plaintext.txt` appear in your folder after a few moments. For more details on all the configuration options for writing to a file you can read [Chapter 5, Section 3](https://macs4200.org/chapters/05/3/file-io.html).

# Submitting your work
You're done with this Lab! All assignments in the course will be distributed as notebooks like this one, and you will submit your work by doing the following:
* Save your notebook
* Restart the kernel and run up to this cell.
* Run all the tests by running the cell containing `grader.check_all()`. Make sure they pass the way you expect them to.
* Run the cell below with the code `grader.export(...)`.
* Download the file named `labXX.zip`, found in the explorer pane on the left side of the screen.
* Upload `labXX-<date-time stamp>.zip` to the corresponding lab assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True, files=['caesar-plaintext.txt'])
```

Tests for q1 (points: 2)
```
>>> text_clean('Hello')
'HELLO'
```
failure_message: text_clean doesn't work with lowercase input text
success_message: text_clean is converting text to upper-case before cleaning
```
>>> text_clean('Hello', match_case=True)
'H'
```
failure_message: text_clean isn't correctly using match_case
success_message: text_clean is correctly using match_case
```
>>> text_clean('Hello, World!')
'HELLOWORLD'
```
failure_message: text_clean is not removing spaces, commas, and exclamation points
success_message: text_clean is removing spaces, commas, and exclamation points
```
>>> text_clean('123ABC!@#def')
'ABCDEF'
```
failure_message: text_clean does not remove punctuation or numerals
success_message: text_clean removes punctuation and numerals
```
>>> text_clean('1234!@#$%^')
''
```
success_message: text_clean returns an empty string if text contains all invalid characters
```
>>> text_clean('')
''
```
success_message: text_clean returns an empty string if text contains an empty string
```
>>> text_clean('a1b2c3', LETTERS='0123456789')
'123'
```
failure_message: text_clean does not consider non-standard LETTERS as inputs
success_message: text_clean can handle various LETTERS as inputs
```
>>> text_clean('hello', LETTERS='HELLO')
'HELLO'
```
```
>>> text_clean('Hello, World!', LETTERS='HLWRD')
'HLLWRLD'
```


Tests for q2 (points: 2)
```
>>> text_block('TESTMESSAGE')
'TESTM ESSAG E'
```
failure_message: Incorrectly uses default n=5 with unequal block size
success_message: Correctly uses default n=5 with unequal block size
```
>>> text_block('abcdefghijklmno', 4)
'abcd efgh ijkl mno'
```
failure_message: Incorrectly uses n=4 with unequal block size
success_message: Correctly uses n=4 with unequal block size
```
>>> text_block('abcdefghi', 3)
'abc def ghi'
```
failure_message: Incorrectly uses n=3 with equal block size
success_message: Correctly uses n=3 with equal block size
```
>>> text_block('abcdefghij', 2)
'ab cd ef gh ij'
```
failure_message: Incorrectly uses n=2 with equal block size
success_message: Correctly uses n=2 with equal block size
```
>>> text_block('a', 3)
'a'
```
failure_message: Does not correctly handle a message shorter than a block
success_message: Correctly handles a message shorter than a block
```
>>> text_block('')
''
```
failure_message: Incorrectly handles an empty message
success_message: Correctly handles an empty message
```
>>> text_block('abcdefgh', 8)
'abcdefgh'
```
failure_message: Incorrectly handles a message exactly equal to a block
success_message: Correctly handles a message exactly equal to a block


Tests for q3 (points: 2)
```
>>> caesar('HELLO', 3) == 'KHOOR'
True
```
failure_message: Basic encryption with default LETTERS
```
>>> caesar('KHOOR', 3, decrypt=True) == 'hello'
True
```
failure_message: Basic decryption with default LETTERS
```
>>> caesar('XYZ', 3) == 'ABC'
True
```
failure_message: Encryption with key wrapping around alphabet
```
>>> caesar('ABC', 3, decrypt=True) == 'xyz'
True
```
failure_message: Decryption with key wrapping around alphabet
```
>>> caesar('Hello, World!', 5) == 'MJQQT BTWQI'
True
```
failure_message: Encrypting a string with spaces and punctuation
```
>>> caesar('MJQQT BTWQI', 5, decrypt=True) == 'helloworld'
True
```
failure_message: Decrypting should return "helloworld" (without spaces)
```
>>> caesar('HELLO', -3) == 'EBIIL'
True
```
failure_message: Encrypting a string with a negative key
```
>>> caesar('EBIIL', -3, decrypt=True) == 'hello'
True
```
failure_message: Decrypting a string with a negative key
```
>>> caesar('HELLO', 0) == 'HELLO'
True
```
failure_message: Encrypting with key = 0 (no shift)
```
>>> caesar('HELLO', 0, decrypt=True) == 'hello'
True
```
failure_message: Decrypting with key = 0 (no shift)
```
>>> caesar('ABCDEFGHIJKLMNOPQRSTUVWXYZ', 3) == 'DEFGH IJKLM NOPQR STUVW XYZAB C'
True
```
failure_message: Encrypting a long string (checks text_block formatting)
```
>>> caesar('DEFGH IJKLM NOPQR STUVW XYZAB C', 3, decrypt=True) == 'abcdefghijklmnopqrstuvwxyz'
True
```
failure_message: Decrypting a long string (checks text_block formatting)


Tests for q4 (points: 2)
```
>>> gcd(48, 18) == 6
True
```
failure_message: gcd(48, 18) should be 6
```
>>> gcd(16, 38) == 2
True
```
failure_message: gcd(16, 38) should be 2
```
>>> gcd(100, 85) == gcd(85, 100)
True
```
failure_message: The GCD function should be commutative (eg. order shouldn't matter)
```
>>> gcd(0, 7) == 7 and gcd(7, 0) == 7
True
```
failure_message: If one number is zero, the GCD should be the value of the other.
```
>>> gcd(5, 5) == 5
True
```
failure_message: When both numbers are the same, the GCD should be that number.
```
>>> gcd(77, 15) == 1
True
```
failure_message: If the numbers are coprime, the GCD should be 1.
```
>>> gcd(8, 32) == 8
True
```
failure_message: gcd shuold work when one number is a multiple of the other


Tests for q5 (points: 1)
```
>>> plaintext == plaintext.lower()
True
```
failure_message: You didn't make the plaintext lowercase
```
>>> len(plaintext)
1154
```
failure_message: Your plaintext isn't the correct length
```
>>> plaintext
'weleftinprettygoodtimeandcameafternightfalltoklausenburghhereistoppedforthenightatthehotelroyaleihadfordinnerorrathersupperachickendoneupsomewaywithredpepperwhichwasverygoodbutthirstymemgetrecipeforminaiaskedthewaiterandhesaiditwascalledpaprikahendlandthatasitwasanationaldishishouldbeabletogetitanywherealongthecarpathiansifoundmysmatteringofgermanveryusefulhereindeedidontknowhowishouldbeabletogetonwithoutithavinghadsometimeatmydisposalwheninlondonihadvisitedthebritishmuseumandmadesearchamongthebooksandmapsinthelibraryregardingtransylvaniaithadstruckmethatsomeforeknowledgeofthecountrycouldhardlyfailtohavesomeimportanceindealingwithanoblemanofthatcountryifindthatthedistricthenamedisintheextremeeastofthecountryjustonthebordersofthreestatestransylvaniamoldaviaandbukovinainthemidstofthecarpathianmountainsoneofthewildestandleastknownportionsofeuropeiwasnotabletolightonanymaporworkgivingtheexactlocalityofthecastledraculaastherearenomapsofthiscountryasyettocomparewithourownordnancesurveymapsbutifoundthatbistritztheposttownnamedbycountdraculaisafairlywellknownplaceishallenterheresomeofmynotesastheymayrefreshmymemorywhenitalkovermytravelswithmina'
```
failure_message: Your plaintext isn't correct
success_message: You correctly decrypted the message


Tests for q6 (points: 1)
```
>>> from pathlib import Path
>>> file_path = Path('caesar-plaintext.txt')
>>> if file_path.exists():
...     text_exists = True
... else:
...     text_exists = False
>>> text_exists
True
```
failure_message: You did not save the file
success_message: You correctly saved the file
